```python
#PLEASE DO NOT MODIFY THIS CELL

import unittest
from unittest.mock import patch
```

# First Semester Project
## Food Ordering Application
A food vendor has approached you build a simple app to help theirs users order food. 
The vendor sells the following food items
* Pizza at 6500 Naira
* Burger at 3000 Naira
* Noodles at 1300 Naira

Write the code for each of these functions to accurately capture a users order and display a receipt to the user.

NOTE: Users can order multiple items in different quantities

## INSTRUCTIONS
* Only type in the areas where you see the comment **"#Your Code Here"**.
* Delete this line **"raise NotImplementedError()"** you will find it directly under the **"#Your code here"** comment.
* Rename the file to the email you used in registering for this course.
* Do not delete any cells or change what is written in them except for the cells that you are supposed to write on.
* This project is worth 30 Marks so take it seriously and adhere strictly to the instructions.

```python
#PLEASE DO NOT MODIFY THIS CELL

def display_menu():
    """
    Description: Prints the menu options for the food items available in the ordering app.
                 since the app only has 3 food items you can order we add a fourth option to exit
                 the menu when they are done ordering.
    """
    print("Menu:")
    print("1. Pizza - 6500")
    print("2. Burger - 3000")
    print("3. Noodles - 1300")
    print("4. Exit Menu")
```

<!-- grade_id: cell-6c24609d20d6ba1a, points: 0, grade: false, solution: true -->
```python
def get_user_choice():
    """
    Description: Takes user input to get the number corresponding to the chosen food item from the menu.
                 Ensures the input is a valid choice between 1 and 4.

                 If the user types a valid number that is not between 1 and 4, return the ERROR
    """
    choice = input("Enter a number for your choice of food ")
    if choice == '4':
        return None
    else:
        pass
    
    if choice not in ['1', '2', '3']:
        return get_user_choice()
        
    else:
        None       
    return int(choice)

```

```python
get_user_choice()
```

Output:
```
Enter a number for your choice of food 1
1
```

<!-- grade_id: cell-3ddef8229d1713a5, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestGetUserChoice(unittest.TestCase):
    @patch('builtins.input', side_effect=['2'])
    def test_valid_choice(self, mock_input):
        result = get_user_choice()
        self.assertEqual(result, 2)

    @patch('builtins.input', side_effect=['invalid', '3'])
    def test_invalid_then_valid_choice(self, mock_input):
        result = get_user_choice()
        self.assertEqual(result, 3)

    @patch('builtins.input', side_effect=['5', '4'])
    def test_invalid_then_exit_choice(self, mock_input):
        result = get_user_choice()
        self.assertIsNone(result)

tester = TestGetUserChoice()
tester.test_valid_choice()
tester.test_invalid_then_valid_choice()
tester.test_invalid_then_exit_choice()
```

<!-- grade_id: cell-c6a618a607e23f46, points: 0, grade: false, solution: true -->
```python
def get_quantity():
    """
    Description: Takes user input to get the quantity of the selected food item.
                 and ensures the input is a positive integer.
    """
    try:
        quantity = int(input("Enter Number of items "))
        if quantity < 0:
            return abs(quantity)
        elif quantity > 0:
            return quantity
        else:
            print('Please enter a value')
            get_quantity() 
    except:
        None
```

```python
get_quantity()
```

Output:
```
Enter Number of items 2
2
```

```python
get_quantity()
```

Output:
```
Enter Number of items -3
3
```

<!-- grade_id: cell-7fd228c806e327c0, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestGetQuantity(unittest.TestCase):
    @patch('builtins.input', side_effect=['3'])
    def test_valid_quantity(self, mock_input):
        result = get_quantity()
        self.assertEqual(result, 3)

    @patch('builtins.input', side_effect=['invalid', '5'])
    def test_invalid_then_valid_quantity(self, mock_input):
        result = get_quantity()
        self.assertEqual(result, 5)

    @patch('builtins.input', side_effect=['0', '-2', '4'])
    def test_invalid_then_valid_quantity_with_negative_input(self, mock_input):
        result = get_quantity()
        self.assertEqual(result, 4)

tester = TestGetQuantity()
tester.test_valid_quantity()
tester.test_invalid_then_valid_quantity()
tester.test_invalid_then_valid_quantity_with_negative_input()
```

Output:
```
File ~\anaconda3\Lib\unittest\case.py:873, in TestCase.assertEqual(self, first, second, msg)
    869 """Fail if the two objects are unequal as determined by the '=='
    870    operator.
    871 """
    872 assertion_func = self._getAssertEqualityFunc(first, second)
--> 873 assertion_func(first, second, msg=msg)

File ~\anaconda3\Lib\unittest\case.py:866, in TestCase._baseAssertEqual(self, first, second, msg)
    864 standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    865 msg = self._formatMessage(msg, standardMsg)
--> 866 raise self.failureException(msg)

AssertionError: None != 5
```

<!-- grade_id: cell-fe851e407b6fa2ce, points: 0, grade: false, solution: true -->
```python
def get_item_name(choice):
    """
    Description: Retrieves and returns the name of a food item based on the user's choice from the menu.
    """
    if choice == 1:
        return 'Pizza'
    elif choice == 2:
        return 'Burger'
    elif choice == 3:
        return 'Noodles'
    elif choice == 4:
        return None
    else:
        "pick between 1 to 4"
        get_item_name(choice)
        
```

```python
get_item_name(1)
```

Output:
```
'Pizza'
```

<!-- grade_id: cell-c061919667620632, points: 3, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

assert get_item_name(1) == 'Pizza'
assert get_item_name(2) == 'Burger'
assert get_item_name(3) == 'Noodles'
```

<!-- grade_id: cell-8266a8ef77b7a512, points: 0, grade: false, solution: true -->
```python
def get_item_price(choice):
    """
    Description: Retrieves and returns the price of a food item based on the user's choice from the menu.
    """
    if choice == 1:
        return 6500
    elif choice == 2:
        return 3000
    elif choice == 3:
        return 1300
    else:
        print("pick between 1 to 3")
        get_item_price(choice)
```

```python
get_item_price(1)
```

Output:
```
6500
```

<!-- grade_id: cell-6b8ba807174ce818, points: 3, grade: true, solution: false -->
```python
assert get_item_price(1) == 6500
assert get_item_price(2) == 3000
assert get_item_price(3) == 1300
```

<!-- grade_id: cell-b5776935997d481b, points: 0, grade: false, solution: true -->
```python
def calculate_total_price(item_price, quantity):
    """
    Description: Calculates and returns the total price of a specific food item based on its price and the quantity ordered.
    """
    total_price = item_price * quantity
    return total_price
```

```python
calculate_total_price(100,2)
```

Output:
```
200
```

<!-- grade_id: cell-2b5e0500570258af, points: 2, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

assert calculate_total_price(5, 2) == 10
```

<!-- grade_id: cell-02a3728e75ef8d90, points: 0, grade: false, solution: true -->
```python
def place_order():
    """
    Description: Manages the process of adding items to a shopping cart. 
                 USE A DICTIONARY FOR THE CART.
                 Calls other functions to get user choices, quantities, and calculates total prices.

                 Your cart should look something like this assuming this user ordered 3 pizzas and 3 burgers.
                 {
                    'Pizza': {'quantity': 3, 'total_price': 19500},
                    'Burger': {'quantity': 3, 'total_price': 9000}
                 }
    """
    
    choice = get_user_choice()
    if choice == None:
        return choice
    else:
        food = get_item_name(choice)
        price = get_item_price(choice)
        num_of_items = get_quantity()
        cart = {
                    food : {'quantity' : num_of_items, 'total_price': calculate_total_price(price, num_of_items)}
                }   

        add_item = input('Do you want to place another order? Y/N ')
        while add_item.upper() == 'Y':
            choice = get_user_choice()
            food = get_item_name(choice)
            price = get_item_price(choice)
            num_of_items = get_quantity()
            cart[food] = {'quantity' : num_of_items, 'total_price': calculate_total_price(price, num_of_items)}
            add_item = input('Do you want to place another order? Y/N ')   
        else:
            pass
    
    
    return cart 
```

```python
place_order()
```

Output:
```
Enter a number for your choice of food 4
```

```python
place_order()
```

Output:
```
Enter a number for your choice of food 1
Enter Number of items 2
Do you want to place another order? Y/N y
Enter a number for your choice of food 2
Enter Number of items 3
Do you want to place another order? Y/N n
{'Pizza': {'quantity': 2, 'total_price': 13000},
 'Burger': {'quantity': 3, 'total_price': 9000}}
```

<!-- grade_id: cell-16918c8d89ee10fa, points: 6, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestPlaceOrder(unittest.TestCase):
    @patch('__main__.get_user_choice', side_effect=[1, 2, None])
    @patch('__main__.get_quantity', return_value=3)
    def test_place_order(self, mock_get_quantity, mock_get_user_choice):
        result = place_order()

        # Assertions based on the expected behavior of place_order
        expected_result = {'Pizza': {'quantity': 3, 'total_price': 19500},
                           'Burger': {'quantity': 3, 'total_price': 9000}}
        self.assertEqual(result, expected_result)

        # Check that get_user_choice was called three times
        self.assertEqual(mock_get_user_choice.call_count, 3)

        # Check that get_quantity was called twice (for the two items added)
        self.assertEqual(mock_get_quantity.call_count, 2)

tester = TestPlaceOrder()
tester.test_place_order()
```

Output:
```
Do you want to place another order? Y/N n
File ~\anaconda3\Lib\unittest\case.py:1172, in TestCase.assertDictEqual(self, d1, d2, msg)
   1168 diff = ('\n' + '\n'.join(difflib.ndiff(
   1169                pprint.pformat(d1).splitlines(),
   1170                pprint.pformat(d2).splitlines())))
   1171 standardMsg = self._truncateMessage(standardMsg, diff)
-> 1172 self.fail(self._formatMessage(msg, standardMsg))

File ~\anaconda3\Lib\unittest\case.py:703, in TestCase.fail(self, msg)
    701 def fail(self, msg=None):
    702     """Fail immediately, with the given message."""
--> 703     raise self.failureException(msg)

AssertionError: {'Piz[28 chars]price': 19500}} != {'Piz[28 chars]price': 19500}, 'Burger': {'quantity': 3, 'total_price': 9000}}
+ {'Burger': {'quantity': 3, 'total_price': 9000},
- {'Pizza': {'quantity': 3, 'total_price': 19500}}
? ^

+  'Pizza': {'quantity': 3, 'total_price': 19500}}
? ^
```

<!-- grade_id: cell-620884e13f967fa3, points: 0, grade: false, solution: true -->
```python
def check_out(cart):
    """
    Description: Finalizes the order by displaying the contents of the shopping c                                   art, including quantities and total prices.
                 Prints the total order price like a receipt.

                 The reciept would look like this if the cart is empty
                 
                     Your cart is empty. No items to check out.


                 If the Cart is has items in it then the receipt should look like this

                     Checking out...
                     Your order details:
                     Item 1: Quantity - 2, Total Price - $20
                     Item 2: Quantity - 3, Total Price - $15
                     Total Order Price: $35
                     Thank you for ordering!
                     
                     
    """
    if cart != None:
        print("\nChecking out...")
        print("Your order details:")
        key_lst = [k for k in cart.keys()]

        index = 0
        total = 0
        while index < len(key_lst):
            food = key_lst[index]
            quantity = cart[key_lst[index]].get('quantity')
            price = cart[key_lst[index]].get('total_price')  
            total = total + price
            index += 1
            print(f"{food}: Quantity : {quantity}, price: ${price}")            
        print(f"Total Order Price: ${total}")
        print("Thank you for ordering!")
    else:
        print("Your cart is empty. No items to check out.")
```

```python
check_out(place_order())
```

Output:
```
Enter a number for your choice of food 4
Your cart is empty. No items to check out.
```

```python
check_out(place_order())
```

Output:
```
Enter a number for your choice of food 1
Enter Number of items 3
Do you want to place another order? Y/N y
Enter a number for your choice of food 2
Enter Number of items 3
Do you want to place another order? Y/N n

Checking out...
Your order details:
Pizza: Quantity : 3, price: $19500
Burger: Quantity : 3, price: $9000
Total Order Price: $28500
Thank you for ordering!
```

<!-- grade_id: cell-db0a762b93d269b1, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestCheckOut(unittest.TestCase):
    @patch('builtins.print')
    def test_check_out_empty_cart(self, mock_print):
        cart = {}
        check_out(cart)
        mock_print.assert_called_with("Your cart is empty. No items to check out.")

    @patch('builtins.print')
    def test_check_out_non_empty_cart(self, mock_print):
        cart = {
            'Item 1': {'quantity': 2, 'total_price': 20},
            'Item 2': {'quantity': 3, 'total_price': 15}
        }
        check_out(cart)

        # Verify that the expected output was printed
        expected_output = [
            "Checking out...",
            "Your order details:",
            "Item 1: Quantity - 2, Total Price - $20",
            "Item 2: Quantity - 3, Total Price - $15",
            "Total Order Price: $35",
            "Thank you for ordering!"
        ]
        calls = [unittest.mock.call(output) for output in expected_output]
        mock_print.assert_has_calls(calls, any_order=False)

tester = TestCheckOut()
tester.test_check_out_empty_cart()
tester.test_check_out_non_empty_cart()
```

Output:
```
Cell In[66], line 8, in TestCheckOut.test_check_out_empty_cart(self, mock_print)
      6 cart = {}
      7 check_out(cart)
----> 8 mock_print.assert_called_with("Your cart is empty. No items to check out.")

File ~\anaconda3\Lib\unittest\mock.py:933, in NonCallableMock.assert_called_with(self, *args, **kwargs)
    931 if actual != expected:
    932     cause = expected if isinstance(expected, Exception) else None
--> 933     raise AssertionError(_error_message()) from cause

AssertionError: expected call not found.
Expected: print('Your cart is empty. No items to check out.')
Actual: print('Thank you for ordering!')
```

```python
#PLEASE DO NOT MODIFY THIS CELL

def food_ordering_app():
    """
    Description: The main function that initiates the food ordering application.
                 Calls place_order() to build the shopping cart and then calls check_out() to complete the order.

                 NOTE THAT IF ANY OF THE OTHER FUNCTIONS ARE NOT CORRECTLY WRITTEN THIS WILL FAIL
                 PLEASE DO NOT MODIFY THIS CELL
    """
    print("Welcome to the Food Ordering App!")
    cart = place_order()
    check_out(cart)
```

```python
food_ordering_app()
```

Output:
```
Welcome to the Food Ordering App!
Enter a number for your choice of food 
Enter a number for your choice of food 1
Enter Number of items 2
Do you want to place another order? Y/N y
Enter a number for your choice of food 2
Enter Number of items 3
Do you want to place another order? Y/N n

Checking out...
Your order details:
Pizza: Quantity : 2, price: $13000
Burger: Quantity : 3, price: $9000
Total Order Price: $22000
Thank you for ordering!
```

<!-- grade_id: cell-bc7e2ab74f6ce4b3, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestFoodOrderingApp(unittest.TestCase):
    @patch('builtins.print')
    @patch('__main__.place_order', return_value={'Item 1': {'quantity': 2, 'total_price': 20}})
    @patch('__main__.check_out')
    def test_food_ordering_app(self, mock_check_out, mock_place_order, mock_print):
        food_ordering_app()

        # Verify that the expected calls were made
        mock_print.assert_called_with("Welcome to the Food Ordering App!")
        mock_place_order.assert_called_once()
        mock_check_out.assert_called_once()

tester = TestFoodOrderingApp()
tester.test_food_ordering_app()
```