# Objects and Data Structures Assessment Test

## Test your knowledge. 

** Answer the following questions **

Write a brief description of all the following Object Types and Data Structures we've learned about: 

* Numbers - Numbers are used to carry out arithmetic calculations. They can either be Integrs (whole numbers eg 18) or Float data type (decimal numbers eg 0.5). 

* Strings are used to create alphanumeric values and are usually in '' or "" For example 'john is a boy' or "john's book is small and he has 2".

* Boolean is used to write logical statements which are either TRUE or False.

* A variable is used as a placeholder for different data types. it is used in conjunction with the assignment operator = . For example surname = 'esheku'.

* A List is a collection of data and can hold different data types, each in their own format. A list is created with square braces []  For example lst = ['john', 34, TRUE]. A list is mutable, it can be manipulated with certain functions like append, insert, delete, pop, extend, remove, sort & sorted etc.

* A Nested list is a list within a list.

* A dictionary is created with curly braces {}. It stores data in a key : Value pair. example: house_names = {'block d' : ['john', 'abu', 'joe']}

* A Tuple is an immutable list it is created by using ().


## Numbers

Write an equation that uses multiplication, division, an exponent, addition, and subtraction that is equal to 100.25.

Hint: This is just to test your memory of the basic arithmetic commands, work backwards from 100.25

```python
# Your answer is probably different
(60 + (10 ** 2) / 4 * 7) - 134.75
```

Output:
```
100.25
```

```python
# my answer

27 * 10 + (9 ** 2) - 1003 / 4
```

Output:
```
100.25
```

Answer these 3 questions without typing code. Then type code to check your answer.

    What is the value of the expression 4 * (6 + 5)
    
    What is the value of the expression 4 * 6 + 5 
    
    What is the value of the expression 4 + 6 * 5 

```python
4 * (6 + 5)
```

* without code BODMAS
 * 4 * (6 + 5)
 * 4 * 11 = 44

```python
# with code

4 * (6 + 5)
```

Output:
```
44
```

```python
4 * 6 + 5 
```

* without code BODMAS
* 4 * 6 + 5 
* 24 + 5 = 29

```python
# with code

4 * 6 + 5 
```

Output:
```
29
```

```python
4 + 6 * 5 
```

* without code BODMAS
* 4 + 6 * 5
* 4 + 30 = 34 

```python
# with code
4 + 6 * 5 
```

Output:
```
34
```

What is the *type* of the result of the expression 3 + 1.5 + 4?

```python
result = 3 + 1.5 + 4

print(result)
```

Output:
```
8.5
```

```python
type(result)
```

Output:
```
float
```

LaTeX:  𝛼2

What would you use to find a number’s square root, as well as its square? 
- Initialize a variable **num_sqrt** to hold the square root of the number
- Initialize a variable **num_sq** to hold the square of the number

```python
num = 100
```

<!-- grade_id: cell-251d4cf6bda089d6, points: 0, grade: false, solution: true -->
```python
# Square root:
num_sqrt = num ** 0.5

print(num_sqrt)
```

Output:
```
10.0
```

<!-- grade_id: cell-af2f50d6feb30fe3, points: 1, grade: true, solution: false -->
```python
assert num_sqrt == 1e1
```

<!-- grade_id: cell-18eeeb77bf75aba5, points: 0, grade: false, solution: true -->
```python
# Square:
num_sq = num ** 2

print(num_sq)
```

Output:
```
10000
```

<!-- grade_id: cell-0032ffc78936cdc3, points: 1, grade: true, solution: false -->
```python
assert num_sq == 1e4
```

## Strings

Given the string 'hello' how would you put the letter 'e' in a variable **letter_e** using the index?.

<!-- grade_id: cell-088358cd80c89256, points: 0, grade: false, solution: true -->
```python
s = 'hello'
letter_e = s[1]

print(letter_e)
```

Output:
```
e
```

<!-- grade_id: cell-8c64ce6723dba25f, points: 1, grade: true, solution: false -->
```python
assert letter_e == 'abcde'[4]
```

Reverse the string 'hello' using slicing and put it in a variable called **reversed**:

<!-- grade_id: cell-de5722380f97a060, points: 0, grade: false, solution: true -->
```python
s ='hello'
# Reverse the string using slicing

reversed = s[::-1]

print(reversed)
```

Output:
```
olleh
```

<!-- grade_id: cell-fc185ccb1afa1ef0, points: 1, grade: true, solution: false -->
```python
assert reversed in 'oiuwaollehimadeel'
```

## Lists

Create a list like [0,0,0] and call it **list_1**.

<!-- grade_id: cell-fac327c594288b28, points: 0, grade: false, solution: true -->
```python
list_1 = [0, 0, 0]

print(list_1)
```

Output:
```
[0, 0, 0]
```

<!-- grade_id: cell-97260ca982a851e2, points: 1, grade: true, solution: false -->
```python
assert list_1 == [0]*3
```

Given **list3** below, change the value 'hello' to 'goodbye'

```python

```

<!-- grade_id: cell-108ce4c1294f7606, points: 0, grade: false, solution: true -->
```python
list3 = [1,2,[3,4,'hello']]

list3[2][2] = 'goodbye'

print(list3)
```

Output:
```
[1, 2, [3, 4, 'goodbye']]
```

<!-- grade_id: cell-2a905f509eb015e6, points: 1, grade: true, solution: false -->
```python
assert list3 == [1, 2, [3, 4, 'goodbye']]
```

Sort the list below and put it in a variable called **sorted_list4**:

<!-- grade_id: cell-986a7f4f32d4209d, points: 0, grade: false, solution: true -->
```python
list4 = [5,3,4,6,1]

sorted_list4 = sorted(list4)

print(sorted_list4)
```

Output:
```
[1, 3, 4, 5, 6]
```

<!-- grade_id: cell-f89627e60e38611d, points: 1, grade: true, solution: false -->
```python
assert sorted_list4 == [1,3,4,5,6]
```

## Dictionaries

Using keys and indexing, grab the 'hello' from the dictionary **d** and put it in a variable called **dict_answer** :

<!-- grade_id: cell-ea72189458dfbb55, points: 0, grade: false, solution: true -->
```python
d = {'k1':[{'nest_key':['this is deep',['hello']]}]}
dict_answer = d['k1'][0]['nest_key'][1][0]

print(dict_answer)
```

Output:
```
hello
```

<!-- grade_id: cell-92c0ac6f156e2a56, points: 1, grade: true, solution: false -->
```python
assert dict_answer in 'asdfheldlorethello8rruhss'

```

## Tuples

Create a tuple called **my_tuple** with the values 1,2 and 3 in it

<!-- grade_id: cell-e778568dc2022c81, points: 0, grade: false, solution: true -->
```python
my_tuple = (1, 2, 3)
```

<!-- grade_id: cell-c84a7e6cfef7d56b, points: 1, grade: true, solution: false -->
```python
assert all(elem in (2,3,5,1,7,8) for elem in my_tuple)
```

## Sets 

Use a set to find the unique values of the list below and assign it to a variable **unique_values**:

<!-- grade_id: cell-9f915e6b6eb8cd39, points: 0, grade: false, solution: true -->
```python
list5 = [1,2,2,33,4,4,11,22,3,3,2]

unique_values = set(list5)
```

<!-- grade_id: cell-9c9f34983f081b5a, points: 1, grade: true, solution: false -->
```python
assert unique_values ==  {1, 2, 3, 4, 11, 22, 33}
```

## Great Job on your first assessment! 